Yay! Exercises! 

```python
import pandas as pd 
from matplotlib import pyplot as plt 
import numpy as np
% matplotlib inline 

from utils import utils 
```

### You're hired! 

Ze and Andreia own a store, and they have just hired you as a data scientist! 

This is going to be fun! 

![Store](https://www.goebt.com/wp-content/uploads/2017/03/iStock-508857266-e1489426963806.jpg)

They have been collecting data over the past few years, and want to start using it to understand their store better. 

Let's get our data: 

```python
store = utils.get_store_data()
store.head()
```

#### Q1: make the index the datetime of the dates

<!-- grade_id: cell-82a9a269d1f23447, points: 0, grade: false, solution: true -->
```python
### Change the index to the corrct date time index. Do this following best practices.  
# Remember, Ze and Andreia are European (hinti-ti-hint-hint)

# store['Date'] = ...
# store = ... 
# store = ...   # best practices! 

store = utils.get_store_data()

# YOUR CODE HERE
raise NotImplementedError()

```

<!-- grade_id: cell-7133328e92b2ffae, points: 2, grade: true, solution: false -->
```python
assert store.index.dtype == '<M8[ns]'
assert store.iloc[25].name == pd.Timestamp('Jan 26th, 2013')

# clue: if this assert is failing, and your iloc[25].name is '04/10/2014', 
# then you are missing the "best practices" part. 
# What did we say in the Learning notebook about this? 
```

#### Q2: what is the most recent day for which we have data? 

<!-- grade_id: cell-697b2952dd12d3eb, points: 0, grade: false, solution: true -->
```python
# get the answer programaticallty 
# hint: the output should be a Timestamp
# most_recent_day = 

# YOUR CODE HERE
raise NotImplementedError()

print(most_recent_day)
```

<!-- grade_id: cell-2ad3e124e392e179, points: 2, grade: true, solution: false -->
```python
assert most_recent_day == pd.Timestamp('July 31st, 2015')
```

#### Q3: How many customers did we have on July 31st, 2015? 

<!-- grade_id: cell-8754da2cad70820d, points: 0, grade: false, solution: true -->
```python
# simple selection, no need to go crazy yet :) 
# number_of_people = ...

# YOUR CODE HERE
raise NotImplementedError()

print(number_of_people)
```

<!-- grade_id: cell-fc854d6f9e221d13, points: 2, grade: true, solution: false -->
```python
assert number_of_people == 555
```

#### Q4: Tuesday, bloody Tuesday 

```python
store.head(2)
```

Yesterday (Thursday) was a very slow day, only 515 people walked in. Andreia wants to know if it's normal. 

What is the median number of people to visit on Thursday?  
- _hint #1: you can use [these methods](https://pandas.pydata.org/pandas-docs/stable/api.html#datetimelike-properties) directly on the index_  
- _hint #2: when operating directly on the index, you do not need to use `.dt` to use the methods_

<!-- grade_id: cell-77921f3ceae2aa39, points: 0, grade: false, solution: true -->
```python
# you will need to do two things. Get the "Thursdays", then take their median. 

# thursday_data = ...
# median_clients_thursday = ... 

# YOUR CODE HERE
raise NotImplementedError()

print(median_clients_thursday)
```

<!-- grade_id: cell-4f2a091d0c170636, points: 2, grade: true, solution: false -->
```python
assert median_clients_thursday == 514.0
```

Ah, maybe it's normal then! 

#### Q5: How was the Christmas season? 

```python
store.head(2)
```

Andreia wants to know how much money they made in total, between the 22nd of December 2014 (inclusive), and the third of January, 2015. 

<!-- grade_id: cell-90c74666589531da, points: 0, grade: false, solution: true -->
```python
# How much money did they make between those two dates? 

# christmas_period_data = ...
# total_money_in_the_christmas_period = ...


# YOUR CODE HERE
raise NotImplementedError()

print(total_money_in_the_christmas_period)
```

<!-- grade_id: cell-665b8b48cc115c92, points: 2, grade: true, solution: false -->
```python
assert total_money_in_the_christmas_period == 53072
```

#### Q6: What was the total weekly revenue, on our worst week of 2014? 

<!-- grade_id: cell-d110e5e35175fdda, points: 0, grade: false, solution: true -->
```python
# Clue: we need weekly data her for this one, 
# Clue: the "worst' week is simply the week when we made the smallest total amount of money

# data_2014 = ...
# total_revenue_per_week_2014 = ...
# worst_week_sales = ...


# YOUR CODE HERE
raise NotImplementedError()

print(worst_week_sales)
```

<!-- grade_id: cell-8b26a15c829c1cba, points: 2, grade: true, solution: false -->
```python
assert worst_week_sales == 14332
```

#### Q7: when was the first time we had a week with more than 4500 clients?  

<!-- grade_id: cell-67c84a8ebd299bd2, points: 0, grade: false, solution: true -->
```python
# note: there are many ways to answer this question, these steps are just to help out. 
# However, we do expect the answer to be a weekly time period (freq='52W-SUN'), so answers  
# which are of the form "the first monday of that week" will not pass the grader. 

# get the days where the store was open 
# store_open_days = 

# make it a weekly dataset, and take the total number of customers 
# total_customers_per_week = ... 

# select only those weeks where we had more than 4500 clients 
# (you can also do this in 2 lines using a mask)
# weeks_when_we_had_more_than_4500_clients = ...

# now get the smallest index of the previous! 
# first_time_we_ever_had_more_than_4500 = ...

# YOUR CODE HERE
raise NotImplementedError()

print(first_time_we_ever_had_more_than_4500)
```

<!-- grade_id: cell-af37d8028da75ec2, points: 2, grade: true, solution: false -->
```python
assert first_time_we_ever_had_more_than_4500 == pd.Timestamp('2013-12-22 00:00:00', freq='52W-SUN')
```

----- 

# Yay! It's an empire now! 

You did such a fantastic job, that Andreia and Ze now own 200 stores. 

You are now their head of data, and have a LOT more data to work with! 

![business empire](https://cdn-images-1.medium.com/max/1280/1*gOmZ5L5WXOKqtw3zWJ_ZGw.jpeg)

```python
data = utils.get_stores_data()
print('We now have %0.0f data points. Wooooow!' % len(data))
```

The thing is, we can't just set the index to be the day, as we now have multiple stores on the same day. 

Looks like we have to go into multi-indexing...

```python
data.head(2)
```

#### Q8: make the date into a datetime, and then set the index to be `[Date, Store]`

<!-- grade_id: cell-bd496ea24bfc5801, points: 0, grade: false, solution: true -->
```python
# make the date a datetime 
# hint: the date format may have changed since we became a big corporate 

data = utils.get_stores_data()

# data.Date = ... 
# YOUR CODE HERE
raise NotImplementedError()

# set the index to be the correct multi-index 
# data = ...   
# data = ...    # don't forget best practices, or it will hurt 
# YOUR CODE HERE
raise NotImplementedError()

print(data.head(3))
```

<!-- grade_id: cell-41fcf00518bc06b3, points: 2, grade: true, solution: false -->
```python
assert isinstance(data.index, pd.core.indexes.multi.MultiIndex)
assert (data.index.levels[0].sort_values()[0:3] == ['2013-01-01', '2013-01-02', '2013-01-03']).all()
assert (data.index.levels[1].sort_values()[0:3] == [1, 2, 3]).all()
```

#### Q9: get the Sales data for store 2 and for store 30, on June 25th 2015

<!-- grade_id: cell-2972eb419a5a4bbf, points: 0, grade: false, solution: true -->
```python
# create a variable called idx, and make it an IndexSlice 
# (clue: if you've forgotten how, check "Selecting in multi-indexes" in the second learning notebook)
# idx = ...

# use the right method to get the answer to Q9. You will need to use idx
# Remember, you want both store 2 and store 30, in one go 
# requested_sales_data = ... 

# YOUR CODE HERE
raise NotImplementedError()

print(requested_sales_data)

```

<!-- grade_id: cell-91c66890cd3120b1, points: 2, grade: true, solution: false -->
```python
assert (requested_sales_data.values == np.array([4022, 3400])).all()
```

#### Q10: Which was the store that had the worse day-to-day absolute fall in sales, excluding days where it was closed? 

<!-- grade_id: cell-fba2c4641714cb4c, points: 0, grade: false, solution: true -->
```python
# This is the champions league question, take some time to think about this one

# A few clues: 
# you will have to do some groupby stuff 
# you will need to get "the index which corresponds to the ...", not the value of the fall
# "the store which had the largest fall"... doesn't this remind you of anything? 
#     (I'll go get an apple from Dingo Poce while you think about it.) 
# when you finally get the index, the store number will be the second element, so you can get it with [1]
# the answer is just the store number.
# printing intermediate results can help a lot 

# store_open = ...
# daily_changes_in_Sales_by_store = ...
# index_of_the_largest_ever_drop = ...
# store_with_the_biggest_drop = ...   # get the element 1 of the tuple 

# YOUR CODE HERE
raise NotImplementedError()

print(store_with_the_biggest_drop)
```

<!-- grade_id: cell-a00120ab691a1e2d, points: 2, grade: true, solution: false -->
```python
assert data.loc[idx['2013-03-06', store_with_the_biggest_drop], 'Customers'].values[0] == 550
```